---
rule_files:
  - ../rules/always-firing.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'vector'
        values: '1+0x100'

    alert_rule_test:
      - eval_time: 3m
        alertname: always-firing
        exp_alerts:
          - exp_annotations:
              summary: This is an always-firing alert to test AlertManager is working
              dashboard: /d/not-implemented
            exp_labels:
              group: mw-monitoring
              receiver: default
              severity: none
