---
rule_files:
  - ../rules/cronjob-failures.yaml

evaluation_interval: 1h

tests:
  - interval: 1h
    input_series:
      - series: 'kube_cronjob_next_schedule_time{job="kube-state-metrics", namespace="prometheus", cronjob="alertmanager-test-cronjob"}'
        values: '0+0x60'
      - series: 'kube_job_spec_completions{job="kube-state-metrics", namespace="prometheus", job_name="alertmanager-test-job-1234"}'
        values: '2+0x60'
      - series: 'kube_job_status_succeeded{job="kube-state-metrics", namespace="prometheus", job_name="alertmanager-test-job-1234"}'
        values: '1+0x60'
      - series: 'kube_job_status_failed{job="kube-state-metrics", namespace="prometheus", job_name="alertmanager-test-job-1234"}'
        values: '1+0x60'

    alert_rule_test:
      - eval_time: 4h     # longer because time() results this
        alertname: long-running-cronjob
        exp_alerts:
          - exp_annotations:
              summary: CronJob prometheus/alertmanager-test-cronjob is taking more than 1h to complete.
              dashboard: /d/not-implemented
            exp_labels:
              group: cronjob-failures
              severity: low
              receiver: slack
              namespace: prometheus
              cronjob: alertmanager-test-cronjob
              job: kube-state-metrics
      - eval_time: 58m
        alertname: long-running-cronjob
        exp_alerts: []
      - eval_time: 2h
        alertname: cronjob-has-not-completed
        exp_alerts:
          - exp_annotations:
              summary: Job prometheus/alertmanager-test-job-1234 is taking more than one hour to complete.
              dashboard: /d/not-implemented
            exp_labels:
              group: cronjob-failures
              severity: low
              receiver: slack
              namespace: prometheus
              job_name: alertmanager-test-job-1234
              job: kube-state-metrics
      - eval_time: 58m
        alertname: cronjob-has-not-completed
        exp_alerts: []
      - eval_time: 2h
        alertname: cronjob-failed
        exp_alerts:
          - exp_annotations:
              summary: Job prometheus/alertmanager-test-job-1234 failed to complete.
              dashboard: /d/not-implemented
            exp_labels:
              group: cronjob-failures
              severity: medium
              receiver: slack
              namespace: prometheus
              job_name: alertmanager-test-job-1234
              job: kube-state-metrics
      - eval_time: 58m
        alertname: cronjob-failed
        exp_alerts: []

