---
rule_files:
  - ../rules/statefulset-failures.yaml

evaluation_interval: 15m

tests:
  - interval: 15m
    input_series:
      - series: 'kube_statefulset_status_observed_generation{job="kube-state-metrics", namespace="prometheus", statefulset="alertmanager-test-statefulset"}'
        values: '2+0x60'
      - series: 'kube_statefulset_metadata_generation{job="kube-state-metrics", namespace="prometheus", statefulset="alertmanager-test-statefulset"}'
        values: '1+0x60'
      - series: 'kube_statefulset_status_replicas{job="kube-state-metrics", namespace="prometheus", statefulset="alertmanager-test-statefulset"}'
        values: '2+0x60'
      - series: 'kube_statefulset_status_replicas_ready{job="kube-state-metrics", namespace="prometheus", statefulset="alertmanager-test-statefulset"}'
        values: '1+0x60'

    alert_rule_test:
      - eval_time: 20m
        alertname: statefulset-generation-mismatch
        exp_alerts:
          - exp_annotations:
              summary: StatefulSet generation for prometheus/alertmanager-test-statefulset does not match, this indicates that the StatefulSet has failed but has not been rolled back.
              dashboard: /d/not-implemented
            exp_labels:
              group: k8s-workloads
              severity: low
              receiver: default
              namespace: prometheus
              statefulset: alertmanager-test-statefulset
              job: kube-state-metrics
      - eval_time: 14m
        alertname: statefulset-generation-mismatch
        exp_alerts: []
      - eval_time: 20m
        alertname: statefulset-replicas-mismatch
        exp_alerts:
          - exp_annotations:
              summary: StatefulSet prometheus/alertmanager-test-statefulset has not matched the expected number of replicas for longer than 15 minutes.
              dashboard: /d/not-implemented
            exp_labels:
              group: k8s-workloads
              severity: moderate
              receiver: default
              namespace: prometheus
              statefulset: alertmanager-test-statefulset
              job: kube-state-metrics
      - eval_time: 14m
        alertname: statefulset-replicas-mismatch
        exp_alerts: []
